---
title: 'Basic Selection'
description: 'This Chapter includes basic column selection techniques that are required'
---

## Insert exercise title here

```yaml
type: NormalExercise
key: bab72e50d8
xp: 100
```



`@instructions`


`@hint`


`@pre_exercise_code`
```{python}

```

`@sample_code`
```{sql}

```

`@solution`
```{sql}

```

`@sct`
```{python}

```
